# Aula 1 - Exercícios - Sintaxe de Python e NumPy Array
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

#### Fonte dos exercícios:
1 - https://www.w3resource.com/python-exercises/python-basic-exercises.php

2 - https://www.practicepython.org/    

### 1.1 Escreva uma função em Python que aceite um inteiro (n) de 1 a 9 e computa o valor de: n+nn+nnn
##### Exemplo: n=5. Então n + nn + nnn = 5 + 55 + 555 = 615

<!-- grade_id: cell-a2f686fdd79e4f51, points: 0, grade: false, solution: true -->
```python
def calculo_n(n):
    """retorne o valor de: n+nn+nnn
    """
    # YOUR CODE HERE
    return n + 10 * n + n + 100 * n + 10 * n + n
    raise NotImplementedError()
```

```python
## Faça seus testes aqui
calculo_n(1)
```

Output:
```
123
```

<!-- grade_id: cell-5412543179b17057, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert calculo_n(1) == 123
assert calculo_n(2) == 246
assert calculo_n(3) == 369
assert calculo_n(9) == 1107
```

### 1.2 Escreva uma função para retornar o maior item de uma lista de inteiros

<!-- grade_id: cell-2e044a5f0541da37, points: 0, grade: false, solution: true -->
```python
# Dica: use float('-inf') para representar um número negativo infinito
def maior(lista):
    """retorne o maior item da lista
    """
    #return max(lista)
    # YOUR CODE HERE
    maior = lista[1]
    for atual in lista : 
        if atual > maior : maior = atual 
    return maior        
    raise NotImplementedError()
```

```python
## Faça seus testes aqui
print (maior([1,0,-1,100,7]))
print (maior([-100,-1580,-1,-9100,-77]))
```

Output:
```
100
-1
```

<!-- grade_id: cell-5bd4079b37b6a7f3, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert maior([1,0,-1,100,7]) == 100
assert maior([-100,-1580,-1,-9100,-77]) == -1
```

### 1.3 Escreva uma função contar a quantidade de strings, presentes numa lista, cujo tamanho é maior que 2 e o primeiro e último caracteres são os mesmos

#### Exemplo:  ['abc', 'xyz', 'aba', '1221']
#### Resultado esperado : 2


<!-- grade_id: cell-5420c3983940f212, points: 0, grade: false, solution: true -->
```python
# Dica 1: A função len retorna o tamanho da string. Exemplo: len('abc'), retorna 3
# Dica 2: Uma string é uma lista de caracteres
def contagem(lista):
    """retorne o maior item da lista
    """
    # YOUR CODE HERE
    qt = 0
    for atual in range(len(lista)) : 
        if len(lista[atual]) > 2 : 
                if lista[atual][:1] == lista[atual][-1] : 
                    qt = qt + 1
    return qt
    
    raise NotImplementedError()
```

```python
## Faça seus testes aqui
print (contagem(['', 'xyzaaaaaaaaa', 'aa', 'babababa']))

```

Output:
```
0
```

<!-- grade_id: cell-f86ec02e2bcb84c6, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert contagem(['abc', 'xyz', 'aba', '1221']) == 2
assert contagem(['ac', '1111', 'aasdfsdfdda', '12ewwww21']) == 3
assert contagem(['', 'xyzaaaaaaaaa', 'aa', 'babababa']) == 0
```

### 1.4 Escreva uma função Python para gerar um dicionário que contenha um número (entre 1 e n) no formato (x, x * x). 
#### Exemplo:  Para n = 5
#### Resultado Esperado: {1: 1, 2: 4, 3: 9, 4: 16, 5: 25}

<!-- grade_id: cell-bb9537795a06bd7e, points: 0, grade: false, solution: true -->
```python
# Dica: use a função range para criar a sequencia de numeros desejados
def create_dict(n):
    """retorne um dicionário com pares (x, x*x)
    """
    # YOUR CODE HERE m = {x: x*x for x in range (1,n+1)}
    dic = {}
    for i in range(1,n + 1):
        dic[i] = i*i
    return dic
    raise NotImplementedError()
```

```python
## Faça seus testes aqui
print (create_dict(5))

```

Output:
```
{1: 1, 2: 4, 3: 9, 4: 16, 5: 25}
```

<!-- grade_id: cell-61d03fd9d1b0eb4e, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert create_dict(5) == {1: 1, 2: 4, 3: 9, 4: 16, 5: 25}

```

### 1.5 Escreva uma função Python para remover itens duplicados de um dicionário, ou seja, cujos valores estejam duplicados.
#### Exemplo:  d =  {1: 1, 2: 1, 3: 9, 4: 1, 5: 25}
#### Resultado Esperado: {1: 1, 3: 9, 5: 25}

<!-- grade_id: cell-1ff4485b8100a753, points: 0, grade: false, solution: true -->
```python
def remove_duplicates(data):
    """retorne um dicionário sem duplicatas
    """
    # YOUR CODE HERE
    result = {}

    for key,value in data.items():
        if value not in result.values():
            result[key] = value
    return result
    raise NotImplementedError()
```

```python
## Faça seus testes aqui
print (remove_duplicates({1: 1, 2: 1, 3: 9, 4: 1, 5: 25}))
data = {'id1': {'name': ['Sara'], 'class': ['V'], 'subject_integration': ['english, math, science']},
        'id2': {'name': ['David'], 'class': ['V'], 'subject_integration': ['english, math, science']},
        'id3': {'name': ['Sara'], 'class': ['V'], 'subject_integration': ['english, math, science']},
        'id4': {'name': ['Surya'], 'class': ['V'], 'subject_integration': ['english, math, science']}
}
print (remove_duplicates(data))
```

Output:
```
{1: 1, 3: 9, 5: 25}
{'id1': {'name': ['Sara'], 'class': ['V'], 'subject_integration': ['english, math, science']}, 'id2': {'name': ['David'], 'class': ['V'], 'subject_integration': ['english, math, science']}, 'id4': {'name': ['Surya'], 'class': ['V'], 'subject_integration': ['english, math, science']}}
```

<!-- grade_id: cell-b6c3a5db7e709b2a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
assert remove_duplicates({1: 1, 2: 1, 3: 9, 4: 1, 5: 25}) ==  {1: 1, 3: 9, 5: 25}
data = {'id1': {'name': ['Sara'], 'class': ['V'], 'subject_integration': ['english, math, science']},
        'id2': {'name': ['David'], 'class': ['V'], 'subject_integration': ['english, math, science']},
        'id3': {'name': ['Sara'], 'class': ['V'], 'subject_integration': ['english, math, science']},
        'id4': {'name': ['Surya'], 'class': ['V'], 'subject_integration': ['english, math, science']}
}
no_duplicates = {'id1': {'name': ['Sara'], 'class': ['V'], 'subject_integration': ['english, math, science']}, 
                 'id2': {'name': ['David'], 'class': ['V'], 'subject_integration': ['english, math, science']}, 
                 'id4': {'name': ['Surya'], 'class': ['V'], 'subject_integration': ['english, math, science']}}
assert remove_duplicates(data) == no_duplicates

```

### 1.6 Escreva uma função Python para combinar os itens de mesma chave de uma lista de  dicionários, somando os valores correspondentes.
##### Exemplo:  para [{'item': 'item1', 'quantidade': 400}, {'item': 'item2', 'quantidade': 300}, {'item': 'item1', 'quantidade': 750}]
##### Resultado esperado: {'item1': 1150, 'item2': 300}

<!-- grade_id: cell-e69a127f71184645, points: 0, grade: false, solution: true -->
```python
def combine_items(item_list):
    """retorne um dicionário combinando (somando) as quantidades do mesmo item
    """
    # YOUR CODE HERE
    result = {}
    for i in range(len(item_list)):
        if item_list[i]['item'] not in result:
            result[item_list[i]['item']] = item_list[i]['quantidade']
        else: result[item_list[i]['item']] = result[item_list[i]['item']] + item_list[i]['quantidade']
    return result
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

input_data = [{'item': 'item1', 'quantidade': 400}, 
              {'item': 'item2', 'quantidade': 300}, 
              {'item': 'item1', 'quantidade': 750}]
output_result = {'item1': 1150, 'item2': 300}
print(combine_items(input_data)) 
```

Output:
```
{'item1': 1150, 'item2': 300}
```

<!-- grade_id: cell-418c099df61d368b, points: 1, grade: true, solution: false -->
```python
# Validação para nota
input_data = [{'item': 'item1', 'quantidade': 400}, 
              {'item': 'item2', 'quantidade': 300}, 
              {'item': 'item1', 'quantidade': 750}]
output_result = {'item1': 1150, 'item2': 300}
assert combine_items(input_data) == output_result

```

### 1.7 Escreva uma função Python para contar a quantidade de itens de um dado valor num array NumPy
#### Exemplo:  para um array([1  4  10  8  4  20  23]) e n=4
#### Resultado esperado: 2
#### Obs.: Lembre-se de usar funções numpy que realizam as operações sem a necessidade de escrever laços (loops)

```python
# import numpy
import numpy as np
```

<!-- grade_id: cell-f6c50b3e165db993, points: 0, grade: false, solution: true -->
```python
# Dica: A função count_nonzero do pacote numpy conta a quantidade 
#       de elementos diferentes de zero ou a quantidade de elementos True.

def count_value(array, n):
    """retorne a quantidade de vezes que n ocorre 
    """
    # YOUR CODE HERE
    return np.count_nonzero(array == n)
    
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

a1 = np.array([1,0,1,9,5,3,1,2,4,5,6,1,2,3,4,8,9,2,1,1,1])
print(count_value(a1, 9))
```

Output:
```
2
```

<!-- grade_id: cell-401910203dadc4db, points: 1, grade: true, solution: false -->
```python
# Validação para nota
a1 = np.array([1,0,1,9,5,3,1,2,4,5,6,1,2,3,4,8,9,2,1])
assert count_value(a1, 1) == 5

```

### 1.8 Escreva uma função Python para somar os valores de cada linha de um array numpy 2D e depois retornar a maior destas somas
#### Exemplo:  para um array([[1  4 ], [10  8], [4  20]]) 
#### Resultado esperado: 24

<!-- grade_id: cell-f7b05dc12adbf2e2, points: 0, grade: false, solution: true -->
```python
# Dica: A função sum do numpy pode fazer o calculo da soma por linha ou por coluna, 
#       desde que o parâmetro axis seja especificado corretamente.
def max_sum(x):
    """retorne o valor maximo dentre os resultados da soma por linha de um array 2D 
    """
    # YOUR CODE HERE
    return max(x.sum(axis=1))
    raise NotImplementedError()
```

```python
## Faça seus testes aqui

a1 = np.array([[1, 4 ], [10, 8], [4, 20]])
a2 = np.array([[1,0,1,9,5],[3,1,2,4,5],[6,1,2,3,4],[8,9,2,1,7]])
print(max_sum(a1)) 
print( max_sum(a2))
```

Output:
```
24
27
```

<!-- grade_id: cell-af1b0435b6aacc9b, points: 1, grade: true, solution: false -->
```python
# Validação para nota
a1 = np.array([[1, 4 ], [10, 8], [4, 20]])
a2 = np.array([[1,0,1,9,5],[3,1,2,4,5],[6,1,2,3,4],[8,9,2,1,7]])
assert max_sum(a1) == 24
assert max_sum(a2) == 27
```

```python

```

### 1.9 Escreva uma função Python para criar um array numpy 2D com 1s nos elementos das bordas e com 0s no interior. O tamanho do array deve ser parametrizável.
#### Exemplo:  para um n=3
#### Resultado esperado: array([[1, 1, 1], [1, 0, 1], [1, 1, 1]])

<!-- grade_id: cell-513f2b547866e072, points: 0, grade: false, solution: true -->
```python
# Dica: A função ones do pacote numpy cria um array de 1s
def create_array(n):
    """retorne um array 2d preenchido com 1s nas bordas e 0s no interior 
    """
    # YOUR CODE HERE
    s = (n,n)
    sz = (n-1,n-1)
    a = np.ones(s)
    a[1:n-1,1:n-1] = 0
    return a
    raise NotImplementedError()
```

```python
## Faça seus testes aqui
print(create_array(2))

```

Output:
```
[[1. 1.]
 [1. 1.]]
```

<!-- grade_id: cell-73baa7b9cc0314dc, points: 1, grade: true, solution: false -->
```python
# Validação para nota
a3 = np.array([[1, 1, 1], [1, 0, 1], [1, 1, 1]])
a6 = np.array([[1., 1., 1., 1., 1., 1.],
       [1., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 1.],
       [1., 1., 1., 1., 1., 1.]])
assert (create_array(3) == a3).all()
assert (create_array(6) == a6).all()


```

### 1.10 Escreva uma função Python para converter para celcius um array de temperaturas em Fahrenheit. Arredonde o resultado até 2a casa decimal usando a função .round(decimals=2) do numpy
#### Exemplo:  para um array([0, 12, 45.21, 34, 99.91])
#### Resultado esperado: array([-17.77, -11.11, 7.33, 1.11, 37.72])

<!-- grade_id: cell-e0a4e5deea9187c4, points: 0, grade: false, solution: true -->
```python
# Dica: T(°C) = (T(°F) - 32) × 5/9
def get_ceusius(F):
    """retorne um array contendo as temperaturas em graus Ceusius
    """
    # YOUR CODE HERE
    C = ((F - 32) * 5 / 9).round(decimals=2)
    return C
    raise NotImplementedError()
```

```python
## Faça seus testes aqui
tf = np.array([0, 12, 45.21, 34, 99.91])
tc = np.array([-17.78, -11.11, 7.34, 1.11, 37.73])
print (get_ceusius(tf)) 

```

Output:
```
[-17.78 -11.11   7.34   1.11  37.73]
```

<!-- grade_id: cell-30d325468b6f6a67, points: 1, grade: true, solution: false -->
```python
# Validação para nota
tf = np.array([0, 12, 45.21, 34, 99.91])
tc = np.array([-17.78, -11.11, 7.34, 1.11, 37.73])
assert (get_ceusius(tf) == tc).all()
```

```python

```