# Aula 3 - Exercícios Extra - Pandas, Fontes de dados e Agregação
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe as bibliotecas pandas, numpy, glob e sqlalchemy
import pandas as pd
import numpy as np

```

### 3.9 Construa uma função para concatenar (verticalmente) o conteúdo de vários arquivos CSV do dataset de morbidade. Esses arquivos contêm dados com colunas compatíveis.
#### Extraia a data do cabeçalho do arquivo. Divida a coluna 'Município' em duas outras contendo o código ibge de 6 caracteres e o nome do município.



<!-- grade_id: cell-c705ab9d11473b9f9, points: 0, grade: false, solution: true -->
```python
# Dica 1: Construa uma lista de arquivos presentes num diretório usando a função glob da biblioteca glob
# Dica 2: Num loop, leia os arquivo duas vezes. Na primeira vez, pra carregar somente a parte desejada do cabeçalho,
#         onde se encontra o mês/ano do respectivo arquivo, utilize o parâmetro nrows. Preste atenção ao separador 
#         dos arquivos CSV, é um ponto e vírgula.
# Dica 3: O método split pode dividir uma string na posição de um caracter especificado.
# Dica 4: Se quiser converter o nome do mês para o seu número, você precisa ajustar a configuração de locale
#         do python para 'pt_BR'. Use o código a seguir para fazer esta operação.
#            import locale
#            locale.setlocale(locale.LC_ALL, 'pt_BR')
# Dica 5: O código strptime(nome_mes,'%b').tm_mon converte nome do mês no seu respectivo número de 1 a 12
# Dica 6: Para criar uma data a partir dos seus atributos use um código semelhante ao exemplo a seguir
#            import datetime
#            datetime.date(year=2020, month=1, day=1)
# Dica 7: Para construir uma nova coluna aplicando uma regra para cada célula use a função apply e construa uma
#         função lambda.
# Dica 8: O método join de uma string consegue unir, com um caracter separador, o conteúdo de uma lista de 
#         elementos numa string. Exemplo: ' '.join(lista_exemplo), produz uma string concatenando com um espaçao
#         os elementos da lista_exemplo. Pode ser útil para isolar o nome do município.
# Dica 9: Adicione todos os dataframes construídos dentro do loop a uma lista e depois use o método concat para 
#         concatenar os dataframes verticalmente. Use o parâmetro ignore_index=True para ignorar o índice dos 
#         dataframes.

def concat_morbidade(files_path):
    """Retorne um dataframe com a concatenação vertical dos dados dos arquivos do diretorio dir_path
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui
files_path = '../../curso_ciencia_dados2020/data/originais/morbidade/2020/*.csv'


```

<!-- grade_id: cell-9ffdsa88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
files_path = '../../curso_ciencia_dados2020/data/originais/morbidade/2020/*.csv'
test_data = {5566: '10899,59', 11082: '16793,57', 16644: '6585,96', 22213: '7446,97', 27779: '19054,46', 33348: '11462,34'}
morbidade = concat_morbidade(files_path)

assert morbidade.loc[morbidade['cod_ibge6']=='522230', 'Valor_total'].to_dict() == test_data
```

### 3.10 Escreva uma função para gerar uma permutação de uma lista e colocar o resultado num dataframe com os elementos da lista agrupados em grupos de tamanho N (Ex.: N=4). 
#### A coluna de índices deve conter o nome dos grupos. Os nomes dos grupos devem seguir a seguinte regra: Grupo 1, Grupo 2, Grupo 3,..., Grupo K.

<!-- grade_id: cell-553qe5248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica: A função len calcula o tamanho de uma lista.

def create_groups(names_list, n):
    """Crie um dataframe com os nomes da lista names_list agrupados em grupos de tamanho n.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui


```

<!-- grade_id: cell-0e2bh6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
import names
N = 9
group_length = 4
test_data = [names.get_full_name() for n in range(N)]
result = ['Grupo 0', 'Grupo 0', 'Grupo 0', 'Grupo 0', 'Grupo 1', 'Grupo 1', 'Grupo 1', 'Grupo 1', 'Grupo 2']
df = create_groups(test_data, group_length)

assert df.index.to_list() == result
```

### 3.11 Escreva uma função para colocar num mesmo dataframe os dados de PIB percapita e índice de gini dos municípios brasieiros. 
#### Retorne os registros completos (PIB percaita, gini, UF, cod_ibge, Municipio) dos top 10 registros com maiores PIB percapita.
#### Analise a diferença entre os quartis (do PIB percapita e do índice de gini) dos top 10 municípios, em termos de PIB percapita,  e os quartis de todo o brasil para estas mesmas duas variáveis. Responda: um PIB percapita significativamente maior, tem ajudado a melhorar o índice de Gini nesses municípios ? Dica: use o método describe(). 
#### Esta análise não vale nota. Será avaliada a nota somente da função que retorna os top 10 registros.

<!-- grade_id: cell-553e5248ebs14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Use as funções criadas no cader de Exercícios desta aula.
# Dica 2: Ao ler o arquivo CSV do índice de gini, se preciso, use encoding='utf8' como parâmetro do método read_csv
def merge_pop_pib(path_pop, path_pib):
    """retorne um dataframe resultante do merge dos dados do pib e da população
    """
    # YOUR CODE HERE
    raise NotImplementedError()

def pib_percapita(pib_pop):
    """retorne retorne um dataframe contendo o codigo ibge de 7 caracteres,
       o nome do municipio e seu PIB percapita.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

def top10(df, col):
    """retorne os top 10 registros do dataframe df ordenados por col
    """
    # YOUR CODE HERE
    raise NotImplementedError()

def top10_municipios(path_pib, path_gini):
    """retorne um dataframe contendo os 10 municipios com maiores PIB percapita. O dataframe deve conter também 
    o valor do índice de gini e nome, UF e código ibge do respectivo município.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui
## Faça seus testes aqui
path_pib = '../../curso_ciencia_dados2020/data/originais/pib/pib_municipios.xlsx'
path_pop='../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'
path_gini='../../curso_ciencia_dados2020/data/originais/gini/ginibr.csv'


# print(df_top10.describe())
# print(pib_gini.describe())

```

<!-- grade_id: cell-0e2bt6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
test_data = {'cod_ibge7': {3677: '3536505', 3547: '3524709'},
 'municipio': {3677: 'Paulínia', 3547: 'Jaguariúna'},
 'uf': {3677: 'SP', 3547: 'SP'},
 'pib_percapita': {3677: 344.8471692406755, 3547: 209.32086192900894},
 'cod_ibge6': {3677: '353650', 3547: '352470'},
 'Município': {3677: 'Paulínia', 3547: 'Jaguariúna'},
 '1991': {3677: '0,49', 3547: '0,5785'},
 '2000': {3677: '0,5734', 3547: '0,5449'},
 '2010': {3677: 0.488, 3547: 0.502}}
assert top10_municipios(path_pib, path_gini).iloc[[0,9]].to_dict() == test_data
```