# Aula 2 - Exercícios - Python e Pandas
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

#### Fonte dos exercícios:
1 - https://www.w3resource.com/python-exercises/python-basic-exercises.php

2 - https://www.kaggle.com/python10pm/pandas-75-exercises-with-solutions

3 - https://www.kaggle.com/python10pm/pandas-exercises-with-solutions-3-3


```python
# Importe as bibliotecas pandas e numpy
import pandas as pd
import numpy as np
```

### 2.1 Escreva uma função para retornar os itens de uma determinada série 
### que não estejam presentes em outra série 
#### Exemplo: Considere uma Serie A com os valores [1, 2, 3, 4, 5, 6, 7, 8], e
####          uma Serie B com os valores [2, 4, 6, 8, 10, 12]
#### Resultado Esperado: Uma Serie com os valores [1, 3, 5, 7]


<!-- grade_id: cell-c705b9d11473b9f9, points: 0, grade: false, solution: true -->
```python
# Dica 1: a função .isin da Serie retorna uma Serie de valores booleanos representando a 
#      intersecção da de duas series
# Dica 2: o operador ~ serve para inverter os valores booleanos de uma Serie.
#      Exemplo: 
#        ~pd.Series([True, False]) 
#        0    False
#        1     True
#        dtype: bool

def complemento(s1, s2):
    """retorne os valores do complemento de s1 em relação a s2
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-9ffda88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
s1 = pd.Series([1, 2, 3, 4, 5, 6, 7, 8]) 
s2 = pd.Series([2, 4, 6, 8, 10, 12]) 
c12 = pd.Series([1, 3, 5, 7]) 
c21 = pd.Series([10, 12]) 

assert complemento(s1,s2).reset_index()[0].equals(c12)
assert complemento(s2,s1).reset_index()[0].equals(c21)
```

### 2.2 Escreva uma função para obter os itens que não são comuns em duas séries.
### Ou seja, estão na uniao mas não na intersecção.
#### Exemplo: Considere uma Serie A com os valores [1, 2, 3, 4, 5, 6, 7, 8], e
####          uma Serie B com os valores [2, 4, 6, 8, 10, 12]
#### Resultado Esperado: Uma Serie com os valores [1, 3, 5, 7,10,12]


<!-- grade_id: cell-553ee248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Use a função np.union1d para calcular a união de duas Series.
#     Exemplo: o código pd.Series(np.union1d(t1, t2)) crie uma Serie da união de t1 e t2

def complemento2(s1, s2):
    """retorne os valores que não são comuns em duas séries
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e2bt6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
s1 = pd.Series([1, 2, 3, 4, 5, 6, 7, 8]) 
s2 = pd.Series([2, 4, 6, 8, 10, 12]) 
c = pd.Series([1, 3, 5, 7, 10, 12]) 

assert complemento2(s1,s2).reset_index()[0].equals(c) 
```

### 2.3 Escreva uma função para selecionar as linhas de um dataframe onde o valor
### na coluna especificada seja ausente (NA)


<!-- grade_id: cell-553d5248eb14bc30, points: 0, grade: false, solution: true -->
```python
def seleciona_null(df, col):
    """retorne as linhas cujo valor na coluna col seja ausente (NA)
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e2b6aefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
exam_data  = {'name': ['Anastasia', 'Dima', 'Katherine', 'James', 'Emily', 'Michael', 'Matthew', 'Laura', 'Kevin', 'Jonas'],
        'score': [12.5, 9, 16.5, np.nan, 9, 20, 14.5, np.nan, 8, 19],
        'attempts': [1, 3, 2, 3, 2, 3, 1, 1, 2, 1],
        'qualify': ['yes', 'no', 'yes', 'no', 'no', 'yes', 'yes', 'no', 'no', 'yes']}
labels = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j']
df = pd.DataFrame(exam_data , index=labels)
assert seleciona_null(df, 'score').equals(df.loc[['d', 'h']])
seleciona_null(df, 'score')
```

### 2.4 Escreva uma função para mudar o conteúdo de uma coluna de um dataframe, 
### especificando como parâmetros de entrada o nome da coluna, o valor original e o 
### novo valor.
#### Exemplo:
|    | name      |   score | qualify   |
|:---|:----------|--------:|:----------|
| a  | Anastasia |    12.5 | yes       |
| b  | Dima      |     9   | no        |
| c  | Katherine |    16.5 | yes       |
| d  | James     |   nan   | no        |

#### Resultado Esperado:
|    | name      |   score | qualify   |
|:---|:----------|--------:|:----------|
| a  | Anastasia |    12.5 | yes       |
| b  | Dima      |     9   | no        |
| c  | Katherine |    16.5 | yes       |
| d  | Suresh    |   nan   | no        |


<!-- grade_id: cell-553f5248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica: Use a função replace da Serie. A função replace tem um parâmetro inplace 
#       para modificar o proprio objeto.

def change_value(df, col, orig, new):
    """altere a coluna col do dataframe df substituindo orig por new
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e2b65cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
exam_data  = {'name': ['Anastasia', 'Dima', 'Katherine', 'James', 'Emily',
                       'Michael', 'Matthew', 'Laura', 'Kevin', 'Jonas'],
        'score': [12.5, 9, 16.5, np.nan, 9, 20, 14.5, np.nan, 8, 19],
        'qualify': ['yes', 'no', 'yes', 'no', 'no', 'yes', 'yes', 'no', 'no', 'yes']}
labels = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j']
exam_data2 = {'name': ['Anastasia', 'Dima', 'Katherine', 'Suresh', 'Emily',
                       'Michael', 'Matthew', 'Laura', 'Kevin', 'Jonas'],
        'score': [12.5, 9, 16.5, np.nan, 9, 20, 14.5, np.nan, 8, 19],
        'qualify': ['yes', 'no', 'yes', 'no', 'no', 'yes', 'yes', 'no', 'no', 'yes']}
df = pd.DataFrame(exam_data, index=labels) 
df2 = pd.DataFrame(exam_data2, index=labels) 

assert change_value(df, 'name', 'James', 'Suresh').equals(df2)
```

### 2.5 Escreva uma função para mudar o conteúdo de uma coluna de um dataframe 
### de valores do tipo string ('yes' ou 'no') para uma coluna do tipo booleano 
### parametrizando no nome da coluna a ser alterada.
#### Exemplo: Considere o dataframe a seguir
|    | name      |   score | qualify   |
|:---|:----------|--------:|:----------|
| a  | Anastasia |    12.5 | yes       |
| b  | Dima      |     9   | no        |
| c  | Katherine |    16.5 | yes       |

#### Resultado Esperado:
|    | name      |   score | qualify   |
|:---|:----------|--------:|:----------|
| a  | Anastasia |    12.5 | True      |
| b  | Dima      |     9   | False     |
| c  | Katherine |    16.5 | True      |


<!-- grade_id: cell-553g5248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica: use a função map da Serie

def change_yes_no(df, col):
    """altere o data frame na coluna col substituindo 'yes' por True e 'no' por False
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e2b6efe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
exam_data  = {'name': ['Anastasia', 'Dima', 'Katherine', 'James', 'Emily',
                       'Michael', 'Matthew', 'Laura', 'Kevin', 'Jonas'],
        'score': [12.5, 9, 16.5, np.nan, 9, 20, 14.5, np.nan, 8, 19],
        'qualify': ['yes', 'no', 'yes', 'no', 'no', 'yes', 'yes', 'no', 'no', 'yes']}
labels = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j']
exam_data2 = {'name': ['Anastasia', 'Dima', 'Katherine', 'James', 'Emily',
                       'Michael', 'Matthew', 'Laura', 'Kevin', 'Jonas'],
        'score': [12.5, 9, 16.5, np.nan, 9, 20, 14.5, np.nan, 8, 19],
        'qualify': [True, False, True, False, False, True, True, False, False, True]}
df = pd.DataFrame(exam_data, index=labels) 
df2 = pd.DataFrame(exam_data2, index=labels) 

assert change_yes_no(df, 'qualify').equals(df2) 
```

### 2.6 Escreva uma função para retornar ordenar um DataFrame 
### por colunas (Ex.: [col1, col2,...,colK] e retornar as top N linhas do DataFrame ordenado 
#### Use o dataset Diamonds (https://www.kaggle.com/shivam2503/diamonds) 
#### O código a seguir faz o download deste dataset.

```python
diamonds = pd.read_csv('https://raw.githubusercontent.com/mwaskom/seaborn-data/master/diamonds.csv')
diamonds
```

<!-- grade_id: cell-553h5248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Use a função sort_values do DataFrame. 
#         Esta função possui os parâmetros inplace e ascending.

def topN_sorted(df, cols, n):
    """retorne as top n linhas ordenadas por cols
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e2b6cere06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
df = pd.DataFrame({'carat':[2.29, 2.0, 1.51],'cut':['Premium', 'Very Good', 'Ideal'],
                   'color':['I', 'G', 'G'],'clarity':['VS2', 'SI1', 'IF'],
                   'depth':[60.8, 63.5, 61.7],'table':[60.0, 56.0, 55.0],
                   'price':[18823, 18818, 18806],'x':[8.5, 7.9, 7.37],
                   'y':[8.47, 7.97, 7.41],'z':[5.16, 5.04, 4.56]},
                 index=[27749, 27748, 27747])

assert topN_sorted(diamonds,['price', 'clarity'], 3).equals(df)
```

### 2.7 Considerando o dataset Diamonds, escreva uma função para retornar as informações
### dos diamantes que atendam a especificação x > N e y > N e Z > N
#### Use o dataset Diamonds (https://www.kaggle.com/shivam2503/diamonds) 
#### O código a seguir faz o download deste dataset


<!-- grade_id: cell-55335248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Use o operador & (AND) para realizar a conjunção lógica (p ^ q).
#         Use a sintaxe com colchete e parênteses df[(S1) & (S2) & (S3)]

def filter_xyz(df, n):
    """retorne as linhas que atendam a especificação x > n e y > n e Z > n
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e24b6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
df = pd.DataFrame({'carat':[2.29, 2.0, 1.51],'cut':['Premium', 'Very Good', 'Ideal'],
                   'color':['I', 'G', 'G'],'clarity':['VS2', 'SI1', 'IF'],
                   'depth':[60.8, 63.5, 61.7],'table':[60.0, 56.0, 55.0],
                   'price':[18823, 18818, 18806],'x':[8.5, 7.9, 7.37],
                   'y':[8.47, 7.97, 7.41],'z':[5.16, 5.04, 4.56]},
                 index=[27749, 27748, 27747])
assert filter_xyz(df, 4).equals(df)
```