# Aula 4 - Exercícios - Agregação com Pandas e Conceitos de Estatística
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

```python
# Importe a biblioteca pandas
import pandas as pd
import numpy as np
```

### 4.1 Escreva uma função para ler os dados de PIB e população dos datasets disponibilizados e retornar um dataframe com a agregação por Estados e as métricas soma, média e mediana do PIB de 2017 e da população de todos os Estados brasileiros.
#### Você pode se basear na sua resolução dos exercícios da Aula 3.


<!-- grade_id: cell-c705ab9d11473b9f9, points: 0, grade: false, solution: true -->
```python
# Dica: Baseie seu código na sua solução do exercício da aula 3 que faz o merge dos
#       datasets do PIB e da população.

def aggregate_pib_pop(path_pop, path_pib):
    """retorne um dataframe com a agregação por Estados e as métricas soma, média, mediana
       e o desvio padrão do PIB de 2017 e da população de todos os Estados brasileiros.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui
path_pib = '../../curso_ciencia_dados2020/data/originais/pib/pib_municipios.xlsx'
path_pop='../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'



```

<!-- grade_id: cell-9ffdsa88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
path_pib = '../../curso_ciencia_dados2020/data/originais/pib/pib_municipios.xlsx'
path_pop='../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'
test_data = {('pop2017', 'sum'): 829619.0,  ('pop2017', 'mean'): 37709.954545454544,
             ('pop2017', 'median'): 17903.5, ('pop2017', 'std'): 79035.97564911412,
             ('2017', 'sum'): 14271061.0, ('2017', 'mean'): 648684.5909090909,
             ('2017', 'median'): 229775.0, ('2017', 'std'): 1695166.9475680576}

assert aggregate_pib_pop(path_pop, path_pib).loc['AC'].to_dict() == test_data
```

### 4.2 Escreva uma função para calcular os percentis [0.25, 0.50, 0.75, 0.9] do dataset do índice de gini dos municípios brasileiros para todos os anos presents no dataset.

#### Apenas por curiosidade, plote o gráfico destes percentis num gráfico de linhas


<!-- grade_id: cell-553qe5248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Descubra qual a string representa valores NA neste dataset.
# Dica 2: Use o argumento na_values para especificar a string que representa valores NA.
# Dica 3: Use o argumento decimal=',' para definir a vírgula como separador de decimal.
# Dica 4: Use a função apply e uma função lambda para desmembrar a coluna município no
#         respectivo nome e código do município.
# Dica 5: Atente para o separador de colunas do dataset.
# Dica 6: O encoding do dataset do gini é 'utf8'.
# Dica 7: Atente para quantas linhas você deve ignorar antes no inicio e no final do arquivo.

def plot_gini(path_gini):
    """retorne o valor de: n+nn+nnn
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui
# Considere o seguinte dataframe do PIB dos municípios
path_gini = '../../curso_ciencia_dados2020/data/originais/gini/ginibr.csv'


```

<!-- grade_id: cell-0e2bh6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
test_data = {'1991': {0.25: 0.48625, 0.5: 0.5314, 0.75: 0.5751, 0.9: 0.617}, 
             '2000': {0.25: 0.5096, 0.5: 0.5535, 0.75: 0.59715, 0.9: 0.6387}, 
             '2010': {0.25: 0.4587, 0.5: 0.5025, 0.75: 0.5459, 0.9: 0.5848}}
path_gini = '../../curso_ciencia_dados2020/data/originais/gini/ginibr.csv'

assert plot_gini(path_gini).to_dict() == test_data
```

### 4.3 Escreva uma função para retornar um dataframe com os nomes das colunas, o total e a porcentagem de valores NA num dataframe passado como argumento para esta função.


<!-- grade_id: cell-553e5248ebs14bc30, points: 0, grade: false, solution: true -->
```python
# Dica: Use a função concat para concatenar duas séries. Especifique axis=1 para concatenar
#       por colunas. Se quiser dar nome as colunas depois de concatenar, utilize o argumento
#       keys=['Total','Porcentage']

def missing_percentage(df):
    """Retorna um dataframe com os nomes das colunas, o total e a porcentagem de valores
       NA no dataframe
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui
path_titanic = '../../curso_ciencia_dados2020/data/originais/titanic/titanic.xls'


def missing_percentage(df):
    """Retorna um dataframe com os nomes das colunas, o total e a porcentagem de valores
       NA no dataframe
    """
    na_count = df.isnull().sum().sort_values(ascending=False)
    na_porce = (100 * df.isnull().sum() / len(df)).sort_values(ascending=False).round(2)
    na_table = pd.concat([na_count, na_porce], axis=1)
    na_table_ren_columns = na_table.rename(columns = {0 : 'Total', 1 : 'Porcentage'})

    return (na_table_ren_columns)
    
    
    df_titanic = pd.read_excel(path_titanic )
    missing_percentage(df_titanic)
```

<!-- grade_id: cell-0e2bt6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
path_titanic = '../../curso_ciencia_dados2020/data/originais/titanic/titanic.xls'
test_data = {'Total': {'body': 1188,  'cabin': 1014,  'boat': 823,  'home.dest': 564,
                       'age': 263,  'embarked': 2,  'fare': 1,  'ticket': 0,  'parch': 0,  
                       'sibsp': 0,  'sex': 0,  'name': 0,  'survived': 0,  'pclass': 0},
             'Porcentage': {'body': 90.76,  'cabin': 77.46,  'boat': 62.87,
                            'home.dest': 43.09,  'age': 20.09,  'embarked': 0.15, 
                            'fare': 0.08,  'ticket': 0.0,  'parch': 0.0,  'sibsp': 0.0,
                            'sex': 0.0,  'name': 0.0,  'survived': 0.0,  'pclass': 0.0}}

df_titanic = pd.read_excel(path_titanic)
assert missing_percentage(df_titanic).to_dict() == test_data
```

### 4.4 Escreva uma função para agregar o dataset do Titanic por classe (pclass) e sexo (sex) e calcular a soma da coluna survived (sobrevivente) e a média e mediana da coluna fare (preço do ticket).


<!-- grade_id: cell-553em5248eb14bc30, points: 0, grade: false, solution: true -->
```python
def agg_titanic(df_titanic):
    """retorne aagregação do dataframe df_titanic por pclass e sex e calcular a soma 
       da coluna survived e a média e mediana da coluna fare.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui
path_titanic = '../../curso_ciencia_dados2020/data/originais/titanic/titanic.xls'


```

<!-- grade_id: cell-0se2b6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
path_titanic = '../../curso_ciencia_dados2020/data/originais/titanic/titanic.xls'
test_data = {('survived', 'sum'): {(1, 'female'): 139,  (1, 'male'): 61,  (2, 'female'): 94,
  (2, 'male'): 25,  (3, 'female'): 106,  (3, 'male'): 75}, ('fare', 'mean'): {(1, 'female'): 109.41238472222216,
  (1, 'male'): 69.88838491620106,  (2, 'female'): 23.234827358490563,  (2, 'male'): 19.90494561403509,
  (3, 'female'): 15.32424953703705,  (3, 'male'): 12.415461991869948}, ('fare', 'median'): {(1, 'female'): 80.92914999999999,
  (1, 'male'): 49.5042,  (2, 'female'): 23.0,  (2, 'male'): 13.0,  (3, 'female'): 10.4896,
  (3, 'male'): 7.8958}}

df_titanic = pd.read_excel(path_titanic)
assert agg_titanic(df_titanic).to_dict() == test_data
```

### 4.5 Escreva uma função para fazer o merge do dataset de morbidade consolidado com o dataset de população
#### O dataset da população deve ficar a esquerda e o dataset da morbidade a direita no resultado.

<!-- grade_id: cell-553a5248eb14bc30, points: 0, grade: false, solution: true -->
```python
def merge_morbidade_pop(path_pop, path_morbidade):
    """retorne o merge dos datasets de população e morbidade.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui
path_morbidade = '../../curso_ciencia_dados2020/data/originais/morbidade/consolidado/morbidade.csv'
path_pop = '../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'



```

<!-- grade_id: cell-0sb2b6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
path_morbidade = '../../curso_ciencia_dados2020/data/originais/morbidade/consolidado/morbidade.csv'
path_pop = '../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'
test_data = {'uf': 'RO', 'cod_uf': '11', 'cod_munic': '00015', 'municipio': "Alta Floresta D'Oeste",
             'pop2017': 25437, 'cod_ibge6': '110001', 'Município': "Alta Floresta D'Oeste",
             'Internações': 149.0, 'AIH_aprovadas': 149, 'Valor_total': 101282.26,
             'Valor_serviços_hospitalares': 77693.11}
df_morbidade_pop = merge_morbidade_pop(path_pop, path_morbidade)
assert df_morbidade_pop.iloc[0, 0:11].to_dict() == test_data

```

### 4.6 Escreva uma função para realizar o agrupamento do dataset de morbidade por UF e por município aplicando a soma sobre as colunas Valor_total, Internações e Dias_permanência. Em seguida calcule, para todos os municípios nesta agregação, o custo médio diário das internações com a seguite fórmula: 
#### custo_medio_diario_intern = Valor_total/Dias_permanência

<!-- grade_id: cell-5b4dc5640fc839be, points: 0, grade: false, solution: true -->
```python
def custo_medio_internacao_agregado(df_morbidade_pop):
    """retorne o merge dos datasets de população e morbidade.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Faça seus testes aqui

```

<!-- grade_id: cell-eee2219340bc517c, points: 1, grade: true, solution: false -->
```python
# Validação para nota
test_data = {'Valor_total': 200333.41999999998, 'Internações': 233.0, 'Dias_permanência': 989.0,
             'custo_medio_diario_intern': 202.56159757330636}
assert custo_medio_internacao_agregado(df_morbidade_pop).iloc[0].to_dict() == test_data
```

### 4.7 Escreva uma função para retornar um dataframe cujos valores de uma dada coluna sejam maiores do que um valor passado como argumento (threshold).
#### Para a validação, utilize o threshold=800. Por curiosidade, plote um boxplot dos dados de custo médio de internação agregados.

<!-- grade_id: cell-919c16309d90227d, points: 0, grade: false, solution: true -->
```python
def get_outliers_morbidade(df_internacoes, col, threshold):
    """retorne um dataframe cujos valores da coluna col sejam maiores que threshold.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Faça seus testes aqui


```

<!-- grade_id: cell-b656ffe7cf513a14, points: 1, grade: true, solution: false -->
```python
# Validação para nota
path_morbidade = '../../curso_ciencia_dados2020/data/originais/morbidade/consolidado/morbidade.csv'
path_pop = '../../curso_ciencia_dados2020/data/originais/populacao/estimativa_dou_2017.xlsx'
col='custo_medio_diario_intern'
threshold=800
test_data = {'Valor_total': 125334.12, 'Internações': 41.0, 'Dias_permanência': 113.0,
             'custo_medio_diario_intern': 1109.1515044247788}

df_morbidade_pop = merge_morbidade_pop(path_pop, path_morbidade)
df_internacoes = custo_medio_internacao_agregado(df_morbidade_pop)

assert get_outliers_morbidade(df_internacoes, col, threshold).iloc[0].to_dict() == test_data
```