# Aula 2 - Exercícios Extra - Python e Pandas
### Consulta a internet (Google, Stackoverflow e etc) liberada!!

#### Fonte dos exercícios:
1 - https://www.w3resource.com/python-exercises/python-basic-exercises.php

2 - https://www.kaggle.com/python10pm/pandas-75-exercises-with-solutions

3 - https://www.kaggle.com/python10pm/pandas-exercises-with-solutions-3-3


```python
# Importe a biblioteca pandas e numpy
import pandas as pd
import numpy as np
```

### 2.8 Escreva uma função para ordenar uma lista de dicionários usando uma função Lambda. 
### A função criada deve receber como argumento a chave do dicionário que será usada 
### na ordenação da lista.


<!-- grade_id: cell-c705b9d11473b9f9, points: 0, grade: false, solution: true -->
```python
# Dica: use a função sorted
def sort_dict_list(l, col):
    """retorne os valores do complemento de s1 em relação a s2
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-9ffda88d58e9282a, points: 1, grade: true, solution: false -->
```python
# Validação para nota
data = [{'name': ['Sara'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']},
        {'name': ['David'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']},
       {'name': ['Surya'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']}]

data_sorted = [{'name': ['David'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']},
        {'name': ['Sara'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']},
       {'name': ['Surya'], 'class': ['V'], 'subject_integration': ['english', 'math', 'science']}]

assert sort_dict_list(data, 'name') == data_sorted
```

### 2.9 Escreva uma função para criar uma nova coluna 'segundo' que tenha 
### o segundo maior valor de cada linha de df.


<!-- grade_id: cell-553ee248eb14bc30, points: 0, grade: false, solution: true -->
```python
# Dica 1: Use a função apply do Dataframe aplicando uma função lambda.
#            e use a função sorted de List dentro da função lambda.
# Dica 2: Aplique a função lambda no eixo 1.

def segundo_maior(df):
    """retorne os valores que não são comuns em duas séries
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e2bt6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
data = {'d1': [20, 36, 25,  4,  1, 20, 17, 47,  2, 40],
        'd2': [33, 25, 48, 39, 49, 37, 45, 36, 37, 31],
        'd3': [19, 47,  9, 27, 48, 26, 24, 21,  1,  2]}
df = pd.DataFrame(data)
assert list(segundo_maior(df)['segundo']) == [20, 36, 25, 27, 48, 26, 24, 36, 2, 31]
```

### 2.10 Escreva uma função para retornar a coluna que contem a maior quantidade de
### valores máximos por linha
#### Exemplo: Considere uma o dataframe a seguir
|    |   d1 |   d2 |   d3 |
|---:|-----:|-----:|-----:|
|  0 |   20 |   33 |   49 |
|  1 |   36 |   25 |   47 |
|  2 |   55 |   48 |    9 |
|  3 |    1 |   12 |    3 |

#### Resultado Esperado: [49, 47,  9, 3]


<!-- grade_id: cell-065ae8d8f787b9b0, points: 0, grade: false, solution: true -->
```python
# Dica 1: Use a função apply do Dataframe aplicando a função np.argmax.
# Dica 2: Aplique a função np.argmax no eixo 1.
# Dica 3: Use a funcao value_counts para contar quantas vezes cada coluna aparece 
#         como índice de máximo valor
# Dica 4: Use o código df.iloc[:,N]  ( onde N é o índice da coluna) para retornar os
#         dados da coluna desejada. Encontre o valor N com a função iloc

def column_of_max_values(df):
    """retorne os valores que não são comuns em duas séries
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
## Faça seus testes aqui

```

<!-- grade_id: cell-0e24b6cefe06bdc60, points: 1, grade: true, solution: false -->
```python
# Validação para nota
df = pd.DataFrame({'d1': [20, 36, 55, 1],
                    'd2': [33, 25, 48, 12],
                    'd3': [49, 47,  9, 3]})
assert list(column_of_max_values(df)) == [49, 47,  9, 3]
```

```python

```